# CPSC 4300/6300-001 Applied Data Science (Fall 2020)

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Joshua Lin"
COLLABORATORS = ""
```

# CPSC4300/6300-001 Problem Set #1

In this assignment, you will create a NCAA College Football team dataset by scraping data from the Internet and perform some basic data manipulations. The assignment has four parts: the first three are auto-graded. Th elast part is manually graded by the TA. 

Before submission, make sure you run all the cells. Because the TA could not run your notebook, your grade will be only based on the output and the writing part of your submitted notebook. 

```python

```

## Part B: Scape the Rosters of Any NCAA Collge Football Team (50 points)

During the class, we have shown how to extract the roster of the Clemson Tiger's defense team, now we want to extend the work to apply to the entire team roster of any NCAA coolege football teams.

### Question 1. Wriet a function to computer the Team URL

Look at the URLs of server college football team at the ESPN site, generalize the pattern, and write a python function  `get_team_url(team_id, team_name)` to compute a url from team_id and team_name. (4 points)

<!-- grade_id: cell-e972833e6447cb14, points: 0, grade: false, solution: true -->
```python
def get_team_url(team_id, team_name):
    """
    Get the team URL for a NCAA college football team
    
    Args:
        team_id: the id of a NACC football team assigned by ESPN
        team_name: the name of the football team
        
    Returns:
        team_url: a string representing the URL of the team
    """
    # YOUR CODE HERE
    #raise NotImplementedError()
    url = "https://www.espn.com/college-football/team/roster/_/id/" + str(team_id) + "/"
    real_team_name = team_name.split( )
    #print(real_team_name)
    i = 1
    while (len(real_team_name) > i):     
        url = url + real_team_name[i-1].lower()+"-"
        i = i + 1
    else: 
        url = url + real_team_name[i-1].lower()
    #print(url)
    return url
```

<!-- grade_id: cell-ce7ed3ddbacff2e2, points: 2, grade: true, solution: false -->
```python
team_id = 228
team_name = "Clemson Tigers"
team_url = get_team_url(team_id, team_name)
assert team_url == 'https://www.espn.com/college-football/team/roster/_/id/228/clemson-tigers'
```

<!-- grade_id: cell-a5c7f35497643616, points: 2, grade: true, solution: false -->
```python
# This cell contains some hidden test cases.
```

### Question 2. Create a BeautifulSoup Object

Complete the following code to create a BeautifulSoup Object from a team url and name the object `soup`. (4 points)

<!-- grade_id: cell-affd9982ef36210b, points: 0, grade: false, solution: true -->
```python
import requests
import re
from bs4 import BeautifulSoup
from requests.exceptions import HTTPError

url = get_team_url(228, "Clemson Tigers")

# YOUR CODE HERE
#raise NotImplementedError()
response = requests.get(url)
html = response.content
soup = BeautifulSoup(html, 'html.parser')
#soup
```

<!-- grade_id: cell-001d464a2c8f9785, points: 2, grade: true, solution: false -->
```python
assert (type(soup) == BeautifulSoup)
```

<!-- grade_id: cell-003e55323653f7bb, points: 2, grade: true, solution: false -->
```python
assert soup.find('title').text.startswith('Clemson')
```

### Question 3. Find the Team Name

Write a statement to find a tag whose text is formatted as "{Team Name} Rosters" on the page. The {Team Name} would be the actual team name like 'Clemson Tigers'. (2 points)

<!-- grade_id: cell-159f19bec0f548a6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()
team_tag = soup.find('h1', string=re.compile(' Roster$'))
```

<!-- grade_id: cell-0313f4fbe0319355, points: 2, grade: true, solution: false -->
```python
assert str(type(team_tag)) == "<class 'bs4.element.Tag'>" and team_tag.text.endswith("Roster")
```

### Question 5. Search for Team Groups

A football team roster is usually divied into three groups: *Offense*, *Defense*, and *Special Teams*. Write some statements to find the tags for all these three groups and save the tags into a variable called groups. (2 points)

<!-- grade_id: cell-7f4a507853f8b142, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()
groups = soup.find_all(class_="Table__Title", string=["Offense", "Defense", "Special Teams"])
print(len(groups), "groups")
for g in groups:
    print(g)
```

Output:
```
3 groups
<div class="Table__Title">Offense</div>
<div class="Table__Title">Defense</div>
<div class="Table__Title">Special Teams</div>
```

<!-- grade_id: cell-42d807cdb2edf390, points: 1, grade: true, solution: false -->
```python
assert ','.join(sorted([g.text for g in groups])) == 'Defense,Offense,Special Teams'
```

<!-- grade_id: cell-71dc4780c401038e, points: 1, grade: true, solution: false -->
```python
assert len(groups) == 3
```

### Question 6. Find the Column Names

Complete the code segment to extract the names of the columns in a table rows and save those names into a list named col_names. (4 points)

<!-- grade_id: cell-2cbb21a90c5b9f64, points: 0, grade: false, solution: true -->
```python
col_names = []
for g in groups:
    t = g.find_next("table")
    for c in t.thead.tr.children:
        col_names.append(c.text)
# YOUR CODE HERE
#raise NotImplementedError()
#col_names
```

<!-- grade_id: cell-ba933afaef916cfd, points: 2, grade: true, solution: false -->
```python
assert "Name" in col_names and "Class" in col_names
```

<!-- grade_id: cell-cf92003b92fab9c5, points: 2, grade: true, solution: false -->
```python
assert type(col_names) is list and len(col_names) >= 6
```

### Question 7. Find the Attributes of a Player

Write code segment to extract the attributes of the first defense player and save them into a dict. The player should have the attributes as shown in the following example.
```
player_name: Lannden Zanders
player_id: 4428971
player_no: 36
POS: S
HT: 6' 1"
WT: 195 lbs
Class: SO
Birthplace: Shelby, NC
```
(10 points)

<!-- grade_id: cell-667e97ed5447158a, points: 0, grade: false, solution: true -->
```python
player = {}
players = []
# YOUR CODE HERE
#raise NotImplementedError()
table = groups[1].find_next("table")
def is_player_row(tag):
    if tag.name == 'tr' and tag.has_attr('data-idx'):
        return True
    return False
for c in table.tbody.find_all(is_player_row):
    tds = c.find_all('td')
    #print(tds)
    cols = []
    for td in tds:
        cols.append(td.get_text())
    players.append(cols)
    
def mysplit(s, ex):
    head = s.rstrip(ex)
    tail = s[len(head):]
    return head, tail

name = mysplit(players[-1][1], '0123456789')
player["player_name"] = name[0]

player_url_name = name[0].split()
totalURL = ""
i = 1
for x in player_url_name :
    if i != len(player_url_name) :
        if x.find(".") :
            x = x.replace('.', '') 
        totalURL = totalURL + x.lower() + "-"
    else :
        if x.find(".") :
            x = x.replace('.', '')
        totalURL = totalURL + x.lower()
    i = i + 1
totalURL = '/' + totalURL

id = soup.find('a', href=re.compile(totalURL))['href']
id = re.sub('http://www.espn.com/college-football/player/_/id/', '', id)
id = re.sub(totalURL, '', id)

player["player_id"] = id
player["player_no"] = name[1]
player["POS"] = players[-1][2]
player["HT"] = players[-1][3]
player["WT"] = players[-1][4]
player["Class"] = players[-1][5]
player["Birthplace"] = players[-1][6]
print(player)
```

Output:
```
{'player_name': 'Lannden Zanders', 'player_id': '4428971', 'player_no': '36', 'POS': 'S', 'HT': '6\' 1"', 'WT': '200 lbs', 'Class': 'SO', 'Birthplace': 'Shelby, NC'}
```

<!-- grade_id: cell-7c574fd6f9413477, points: 1, grade: true, solution: false -->
```python
assert type(player) is dict
```

<!-- grade_id: cell-75b251071e223227, points: 2, grade: true, solution: false -->
```python
assert len(player) == 8
```

<!-- grade_id: cell-9c5f3fdbcd88b09b, points: 5, grade: true, solution: false -->
```python
assert ','.join(sorted(player.keys())) == 'Birthplace,Class,HT,POS,WT,player_id,player_name,player_no'
```

<!-- grade_id: cell-7071bacb49ae8bef, points: 1, grade: true, solution: false -->
```python
assert int(player['player_no']) > 0 and int(player['player_no']) < 200
```

<!-- grade_id: cell-232a559e5c4360cc, points: 1, grade: true, solution: false -->
```python
assert int(player['WT'].split()[0]) > 150 and player['WT'].split()[1] == 'lbs'
```

### Question 8. Abstract the tasks into a function

Since you don't want to repeat the above steps for every player and every team, it is necessary to abstract the above scraping task into a function which you can reuse over and over. (10 points)

<!-- grade_id: cell-3de336d8b0a64126, points: 0, grade: false, solution: true -->
```python
import requests
import re
from requests.exceptions import HTTPError
from bs4 import BeautifulSoup

def extract_team_roster(team_id, team_name):
    """
    Extract the team roster for a NCAA college football team
    
    Args:
        team_id: the id of a NACC football team assigned by ESPN
        team_name: the name of the football team
        
    Returns:
        a list of dict representing the roster of the entire team.
    """
    print("wrong")
    #team_id = 150
    #team_name = "Duke Blue Devils"
    roster = []
    def mysplit(s, ex):
        head = s.rstrip(ex)
        tail = s[len(head):]
        return head, tail
    
    def is_player_row(tag):
        if tag.name == 'tr' and tag.has_attr('data-idx'):
            return True
        return False
    
    #Computing Team URL
    url = "https://www.espn.com/college-football/team/roster/_/id/" + str(team_id) + "/"
    real_team_name = team_name.split( )
    #print(real_team_name)
    i = 1
    while (len(real_team_name) > i):     
        url = url + real_team_name[i-1].lower()+"-"
        i = i + 1
    else: 
        url = url + real_team_name[i-1].lower()

    #Create a BeautifulSoup Object
    response = requests.get(url)
    html = response.content
    soup = BeautifulSoup(html, 'html.parser')
        
    #Finding Team Name
    team_tag = soup.find('h1', string=re.compile(' Roster$'))
        
    #Searching For Team Group
    groups = soup.find_all(class_="Table__Title", string=["Offense", "Defense", "Special Teams"])
    #print(len(groups), "groups")
    #for g in groups:
        #print(g)
        
    #Finding the Column Names
    col_names = []
    for g in groups:
        t = g.find_next("table")
        for c in t.thead.tr.children:
            col_names.append(c.text)
        
    #Find the Attributes of a Player
    players = []
    id_list = []
    for x in groups :
        table = x.find_next("table")
        for c in table.tbody.find_all(is_player_row):
            tds = c.find_all('td')
            cols = []
            for td in tds:
                cols.append(td.get_text())
            players.append(cols)
        for aPlayer in players :
            player = {}
            name = mysplit(aPlayer[1], '0123456789')
            player["player_name"] = name[0]

            player_url_name = name[0].split()
            totalURL = ""
            i = 1
            for x in player_url_name :
                if i != len(player_url_name) :
                    if x.find(".") :
                        x = x.replace('.', '') 
                    if x.find('\'') :
                        x = x.replace('\'', '') 
                    totalURL = totalURL + x.lower() + "-"
                else :
                    if x.find(".") :
                        x = x.replace('.', '')
                    if x.find('\'') :
                        x = x.replace('\'', '') 
                    totalURL = totalURL + x.lower()
                i = i + 1
            totalURL = '/' + totalURL
            #print(totalURL)
            #print(soup.find('a', href=re.compile(totalURL))['href'])
            id = soup.find('a', href=re.compile(totalURL))['href']
            id = re.sub('http://www.espn.com/college-football/player/_/id/', '', id)
            id = re.sub(totalURL, '', id)
            if id_list.count(id) == 0: 
                id_list.append(id)
                player["player_id"] = id
                if name[1] == '' :
                    player["player_no"] = "NA"
                else :
                    player["player_no"] = name[1]
                player["POS"] = aPlayer[2]
                player["HT"] = aPlayer[3]
                player["WT"] = aPlayer[4]
                player["Class"] = aPlayer[5]
                player["Birthplace"] = aPlayer[6]
                #print(player)
                roster.append(player)
        
    # YOUR CODE HERE
    #raise NotImplementedError()
    return roster
```

```python
import pprint
import numpy as np
team_name = "Virginia Cavaliers"
team_id = 258
team_roster = extract_team_roster(team_id, team_name)
pp = pprint.PrettyPrinter(indent=4)
#pp.pprint(team_roster)
player_ids = np.array([p['player_id'] for p in team_roster])
unqiue_ids = np.unique(player_ids)
#assert len(unqiue_ids) == len(team_roster)
#print(len(unqiue_ids))
#print(len(team_roster))
player_nos = np.array([p['player_no'] for p in team_roster])
na_count = len([no for no in player_nos if no == 'NA'])
print(na_count)
#assert na_count > 10
```

Output:
```
wrong
45
```

```python
# test the function
import pprint

team_id = 150
team_name = "Duke Blue Devils"
team_roster = extract_team_roster(team_id, team_name)

pp = pprint.PrettyPrinter(indent=4)
pp.pprint(team_roster[:3])
```

Output:
```
wrong
[   {   'Birthplace': 'Grayson, GA',
        'Class': 'JR',
        'HT': '6\' 3"',
        'POS': 'QB',
        'WT': '235 lbs',
        'player_id': '4239994',
        'player_name': 'Chase Brice',
        'player_no': '8'},
    {   'Birthplace': 'Langhorne, PA',
        'Class': 'SO',
        'HT': '6\' 3"',
        'POS': 'QB',
        'WT': '190 lbs',
        'player_id': '4568231',
        'player_name': 'Jack Colyar',
        'player_no': '13'},
    {   'Birthplace': 'Venice, CA',
        'Class': 'FR',
        'HT': '6\' 3"',
        'POS': 'QB',
        'WT': '180 lbs',
        'player_id': '4429578',
        'player_name': 'Luca Diamont',
        'player_no': '2'}]
```

<!-- grade_id: cell-895004dba91372be, points: 1, grade: true, solution: false -->
```python
type(team_roster) is list
```

Output:
```
True
```

<!-- grade_id: cell-018fc6188aad6ea2, points: 1, grade: true, solution: false -->
```python
assert len(team_roster) >= 80
```

<!-- grade_id: cell-a5e0daeb322a2166, points: 4, grade: true, solution: false -->
```python
assert ','.join(sorted(team_roster[50].keys())) == 'Birthplace,Class,HT,POS,WT,player_id,player_name,player_no'
```

<!-- grade_id: cell-fdf01ad39dd2e479, points: 2, grade: true, solution: false -->
```python
import numpy as np
classes = np.array([p['Class'] for p in team_roster])
class_labels, class_counts = np.unique(classes, return_counts=True)
assert ''.join(class_labels.tolist()) == 'FRJRSOSR'
```

<!-- grade_id: cell-d6f3d4de23bc3c91, points: 2, grade: true, solution: false -->
```python
import numpy as np
player_ids = np.array([p['player_id'] for p in team_roster])
unqiue_ids = np.unique(player_ids)
assert len(unqiue_ids) == len(team_roster)
```

### Question 9. Debug and Testing Functions

When you write a program, you have to make sure that it always works correctly even if some exteral data (e.g., the ESPN website) may be corrupted. For example, Some players in a certain team may not have their player no. Your code should be able tp catch such exceptions and replace the missing player no with the value "NA". (10 points)

<!-- grade_id: cell-87356ddb0aaeca8e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()
import requests
import re
from requests.exceptions import HTTPError
from bs4 import BeautifulSoup
team_name = "Virginia Cavaliers"
team_id = 258

print("wrong3")
roster = []
def mysplit(s, ex):
    head = s.rstrip(ex)
    tail = s[len(head):]
    return head, tail

def is_player_row(tag):
    if tag.name == 'tr' and tag.has_attr('data-idx'):
        return True
    return False

#Computing Team URL
url = "https://www.espn.com/college-football/team/roster/_/id/" + str(team_id) + "/"
real_team_name = team_name.split( )
#print(real_team_name)
i = 1
while (len(real_team_name) > i):     
    url = url + real_team_name[i-1].lower()+"-"
    i = i + 1
else: 
    url = url + real_team_name[i-1].lower()

#Create a BeautifulSoup Object
response = requests.get(url)
html = response.content
soup = BeautifulSoup(html, 'html.parser')

#Finding Team Name
team_tag = soup.find('h1', string=re.compile(' Roster$'))

#Searching For Team Group
groups = soup.find_all(class_="Table__Title", string=["Offense", "Defense", "Special Teams"])
#print(len(groups), "groups")
#for g in groups:
    #print(g)

#Finding the Column Names
col_names = []
for g in groups:
    t = g.find_next("table")
    for c in t.thead.tr.children:
        col_names.append(c.text)

#Find the Attributes of a Player
players = []
id_list = []
for x in groups :
    table = x.find_next("table")
    for c in table.tbody.find_all(is_player_row):
        tds = c.find_all('td')
        cols = []
        for td in tds:
            cols.append(td.get_text())
        players.append(cols)
    for aPlayer in players :
        player = {}
        name = mysplit(aPlayer[1], '0123456789')
        player["player_name"] = name[0]

        player_url_name = name[0].split()
        totalURL = ""
        i = 1
        for x in player_url_name :
            if i != len(player_url_name) :
                if x.find(".") :
                    x = x.replace('.', '') 
                if x.find('\'') :
                    x = x.replace('\'', '') 
                totalURL = totalURL + x.lower() + "-"
            else :
                if x.find(".") :
                    x = x.replace('.', '')
                if x.find('\'') :
                    x = x.replace('\'', '') 
                totalURL = totalURL + x.lower()
            i = i + 1
        totalURL = '/' + totalURL
        #print(totalURL)
        #print(soup.find('a', href=re.compile(totalURL))['href'])
        id = soup.find('a', href=re.compile(totalURL))['href']
        id = re.sub('http://www.espn.com/college-football/player/_/id/', '', id)
        id = re.sub(totalURL, '', id)
        if id_list.count(id) == 0: 
            id_list.append(id)
            player["player_id"] = id
            if name[1] == '' :
                player["player_no"] = "NA"
            else :
                player["player_no"] = name[1]
            player["POS"] = aPlayer[2]
            player["HT"] = aPlayer[3]
            player["WT"] = aPlayer[4]
            player["Class"] = aPlayer[5]
            player["Birthplace"] = aPlayer[6]
            #print(player)
            roster.append(player)

team_roster = roster
```

Output:
```
wrong3
```

<!-- grade_id: cell-0ed1c994a0f7210c, points: 5, grade: true, solution: false -->
```python
import numpy as np
positions = np.array([p['POS'] for p in team_roster])
position_label, position_counts = np.unique(positions, return_counts=True)
assert ''.join(position_label.tolist()) == 'CCBDBDEDLDTFBGLBLSNTOLOTPPKQBRBSTEWR'
```

<!-- grade_id: cell-cf43e6715aff64de, points: 5, grade: true, solution: false -->
```python
player_nos = np.array([p['player_no'] for p in team_roster])
na_count = len([no for no in player_nos if no == 'NA'])
assert na_count > 10
```

### Question 10. Add functions into a Python Modules

Now you have developed two solid functions: `get_team_url` and `extract_team_roster` and you would like to reuse them in other projects or share them with other people. Use the Notebook magic command to write the two functions with other neccessary  statements into a python file named `roster.py`. (4 points)

<!-- grade_id: cell-9121f8f48dd11851, points: 0, grade: false, solution: true -->
```python
%%writefile roster.py

# YOUR CODE HERE
#raise NotImplementedError()
import requests
import re
from requests.exceptions import HTTPError
from bs4 import BeautifulSoup

def get_team_url(team_id, team_name):
    url = "https://www.espn.com/college-football/team/roster/_/id/" + str(team_id) + "/"
    real_team_name = team_name.split( )
    #print(real_team_name)
    i = 1
    while (len(real_team_name) > i):     
        url = url + real_team_name[i-1].lower()+"-"
        i = i + 1
    else: 
        url = url + real_team_name[i-1].lower()
    return url

def extract_team_roster(team_id, team_name):
    """
    Extract the team roster for a NCAA college football team
    
    Args:
        team_id: the id of a NACC football team assigned by ESPN
        team_name: the name of the football team
        
    Returns:
        a list of dict representing the roster of the entire team.
    """
    #team_id = 150
    #team_name = "Duke Blue Devils"
    roster = []
    def mysplit(s, ex):
        head = s.rstrip(ex)
        tail = s[len(head):]
        return head, tail
    
    def is_player_row(tag):
        if tag.name == 'tr' and tag.has_attr('data-idx'):
            return True
        return False
    
    #Computing Team URL
    url = get_team_url(team_id, team_name)

    #Create a BeautifulSoup Object
    response = requests.get(url)
    html = response.content
    soup = BeautifulSoup(html, 'html.parser')
        
    #Finding Team Name
    team_tag = soup.find('h1', string=re.compile(' Roster$'))
        
    #Searching For Team Group
    groups = soup.find_all(class_="Table__Title", string=["Offense", "Defense", "Special Teams"])
    #print(len(groups), "groups")
    #for g in groups:
        #print(g)
        
    #Finding the Column Names
    col_names = []
    for g in groups:
        t = g.find_next("table")
        for c in t.thead.tr.children:
            col_names.append(c.text)
        
    #Find the Attributes of a Player
    players = []
    id_list = []
    for x in groups :
        table = x.find_next("table")
        for c in table.tbody.find_all(is_player_row):
            tds = c.find_all('td')
            cols = []
            for td in tds:
                cols.append(td.get_text())
            players.append(cols)
        for aPlayer in players :
            player = {}
            name = mysplit(aPlayer[1], '0123456789')
            player["player_name"] = name[0]
            
            player_url_name = name[0].split()
            #print(player_url_name)
            totalURL = ""
            i = 1
            for x in player_url_name :
                #print(x)
                if i != len(player_url_name) :
                    if x.find(".") :
                        x = x.replace('.', '') 
                    if x.find('\'') :
                        #print(x)
                        x = x.replace('\'', '') 
                    if x.find(',') :
                        x = x.replace(',', '')
                    #print(x)
                    totalURL = totalURL + x.lower() + "-"
                else :
                    if x.find(".") :
                        x = x.replace('.', '')
                    if x.find('\'') :
                        #print(x)
                        x = x.replace('\'', '') 
                    if x.find(',') :
                        x = x.replace(',', '')
                    if x != '' :
                        #print(x)
                        totalURL = totalURL + x.lower()
                i = i + 1
            totalURL = '/' + totalURL
            #print(totalURL)
            #print(soup.find('a', href=re.compile(totalURL))['href'])
            id = soup.find('a', href=re.compile(totalURL))['href']
            id = re.sub('http://www.espn.com/college-football/player/_/id/', '', id)
            id = re.sub(totalURL, '', id)
            if id_list.count(id) == 0: 
                id_list.append(id)
                player["player_id"] = id
                if name[1] == '' :
                    player["player_no"] = "NA"
                else :
                    player["player_no"] = name[1]
                player["POS"] = aPlayer[2]
                player["HT"] = aPlayer[3]
                player["WT"] = aPlayer[4]
                player["Class"] = aPlayer[5]
                player["Birthplace"] = aPlayer[6]
                #print(player)
                roster.append(player)
        
    # YOUR CODE HERE
    #raise NotImplementedError()
    return roster
```

Output:
```
Overwriting roster.py
```

```python
#import os
#import json
#import pandas as pd
#import numpy as np
#from roster import extract_team_roster
#hi = extract_team_roster(2390, 'Miami Hurricanes')
#print(hi)
```

```python
!cat roster.py
```

Output:
```
# YOUR CODE HERE
#raise NotImplementedError()
import requests
import re
from requests.exceptions import HTTPError
from bs4 import BeautifulSoup

def get_team_url(team_id, team_name):
    url = "https://www.espn.com/college-football/team/roster/_/id/" + str(team_id) + "/"
    real_team_name = team_name.split( )
    #print(real_team_name)
    i = 1
    while (len(real_team_name) > i):     
        url = url + real_team_name[i-1].lower()+"-"
        i = i + 1
    else: 
        url = url + real_team_name[i-1].lower()
    return url

def extract_team_roster(team_id, team_name):
    """
    Extract the team roster for a NCAA college football team
    
    Args:
        team_id: the id of a NACC football team assigned by ESPN
        team_name: the name of the football team
        
    Returns:
        a list of dict representing the roster of the entire team.
    """
    #team_id = 150
    #team_name = "Duke Blue Devils"
    roster = []
    def mysplit(s, e
```

<!-- grade_id: cell-460dca90f00e20e6, points: 1, grade: true, solution: false -->
```python
import os
assert os.path.exists("roster.py")
```

<!-- grade_id: cell-d2d114880fe61a22, points: 3, grade: true, solution: false -->
```python
try:
    from bs4 import BeautifulSoup
    import requests
    from roster import get_team_url, extract_team_roster
    
    clemson_tiger_roster = extract_team_roster(258, "Virginia Cavaliers")
except:
    raise AssertionError("could not import")

```

### Congratulations! Yo have finished Part B of the problem.